---
title: 'Python Basis'
description: 'In de eerste les leer je hoe je de print functie werkt.'
---

## Je eerste python code.

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

In deze eenvoudige oefening ga je je eerste Python code schrijven. Deze tekst zegt niet veel. Hierbij zal je enkel en **alleen** werken met de functie ```
print 
```
De rest is history.

`@instructions`
Print de tekst met de volgende woorden: **Hello World.**

`@hint`
Gewoon blijven proberen tot het je lukt!

`@pre_exercise_code`
```{python}


```

`@sample_code`
```{python}
#Print de tekst
```

`@solution`
```{python}
Het antwoord is print("Hello World")
```

`@sct`
```{python}
success_msg("Goed gedaan")
```
